<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

Parcours séquentiel de listes
======
Codage des algorithmes de parcours en Python
------


# Préambule
L'objectif de cette partie est de comprendre comment fonctionnent ces algorithmes de parcours, pour cela il conviendra de n'utiliser que des opérations de boucle (`while` ou `for`) et des opérations conditonnelles (`if`).

```python
liste = "a,b,c"
x = []
for i in liste:
    if i != ",":
        x.append(i)
x

v = [i for i in liste if i!= ","]
v
```

Output:
```
['a', 'b', 'c']
```

# Recherche d'un maximum
## Recherche d'une valeur maximale
> ecrire une fonction `maximum(T)` qui renvoie la valeur maximale d'une liste d'entiers.

<!-- grade_id: cell-7b918e9bde49dd86, points: 0, grade: false, solution: true -->
```python
def maximum(J):
    max_value = J[0]
    for x in range(len(J)):
        while max_value < J[x]:
            max_value = J[x]
    return max_value
```

<!-- grade_id: cell-1ec3eaf5cc260f21, points: 2, grade: true, solution: false -->
```python
assert maximum([1,2,3,4,5,6,7,8,9,10]) == 10
assert maximum([1]) == 1
assert maximum([10,9,8,7,6,5,4,3,2,1,0]) == 10
assert maximum([10,9,8,7,6,100,5,4,3,2,1,0]) == 100
```

## Complexité de la recherche d'une valeur maximale
Nous allons maintenant vérifier que le coût en temps d'un tel algorithme est linéaire. Considérons la fonction `compte_temps(n)`qui renvoie le temps mis la fonction `maximum(T)` pour s'éxécuter.

```python
from time import process_time
from random import randint
def compte_temps(n):
    L = [randint(1,100) for i in range(n) ]
    t1_start = process_time()  
    maximum(L)
    t1_stop = process_time() 
    return (t1_stop-t1_start)  
```

```python
import matplotlib.pyplot as plt

x = [i for i in range(0,100000,5000)]
y=[]
for i in x :
    y.append(compte_temps(i))

plt.plot(x,y)
```

**Quelle est la grandeur en abcisse ? la grandeur en ordonnée ?**

<!-- grade_id: cell-14a24e9e1f79ba8b, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**A partir du graphe obtenu justifiez que la complexité de l'algorithme de recherche d'un maximum est linéaire.**

<!-- grade_id: cell-13e25a8bd5da3148, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Estimez le temps que mettrait votre algorithme à effectuer la recherche d'un maximum dans une liste comportant $10^9$ éléments.**

<!-- grade_id: cell-30d59d4deeef97c5, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Recherche de l'indice d'une valeur maximale
> Ecrire une fonction `indice_max(T)` qui renvoie l'indice du maximum d'une liste d'entiers.

<!-- grade_id: cell-334df6fc3ecafca8, points: 0, grade: false, solution: true -->
```python
def indice_max(T):
    for i in range(len(T)):
        if T[i] == maximum(T):
            return i
```

<!-- grade_id: cell-dbc84904811bad1f, points: 2, grade: true, solution: false -->
```python
assert indice_max([1,2,3,4,5]) == 4
assert indice_max([6,2,3,4,5,3]) == 0
assert indice_max([1,2,3,5,4,3,2,1]) == 3
### BEGIN HIDDENT TESTS
assert indice_max([1]) == 0
assert indice_max([0,1,1,1,1,1,1]) == 1
```

Nous allons maintenant vérifier que le cout de cet algorithme `indice_max(T)`de recherche est bien lui aussi linéaire, mais aussi nous allons le comparer à celui de la fonction précédente `maximum(T)`

```python
def compte_temps2(n):
    L = [randint(1,100) for i in range(n) ]
    t1_start = process_time()  
    indice_max(L)
    t1_stop = process_time() 
    return (t1_stop-t1_start) 
```

```python
import matplotlib.pyplot as plt

x = [i for i in range(0,100000,5000)]
y1,y2=[],[]
for i in x :
    y1.append(compte_temps(i))
    y2.append(compte_temps2(i))
plt.plot(x,y1,label="valeur")
plt.plot(x,y2,label="indice")
plt.legend()
```

**Vérifier sur le graphe ci-dessus que les complexités des deux algorithmes sont bien linéaires**

<!-- grade_id: cell-2bf2ce715bc918f1, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Expliquer la différence de temps d'éxécution entre les fonctions `maximum(T)` et `indice_max(T)`**

<!-- grade_id: cell-a7e00ac9fc9001c3, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE