<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

Parcours séquentiel de listes
======
Codage des algorithmes de parcours en Python
------


# Préambule
L'objectif de cette partie est de comprendre comment fonctionnent ces algorithmes de parcours, pour cela il conviendra de n'utiliser que des opérations de boucle (`while` ou `for`) et des opérations conditonnelles (`if`). **Il ne sera pas possible par de coder `if i in T` attention !!!**

# Algorithmes de vérification
## Test d'appartenance
> Ecrire une fonction `appartient(T,n)` qui teste l'appartenance de `n` à la liste `T`

**Attetntion vous n'avez pas le droit d'utiliser la strucuture `if i in T : return True`pour écrire votre fonction !**

Appuyer vous sur l'algorithme vu en cours avec une boucle `while..`

<!-- grade_id: cell-b5ffb1463aaf1421, points: 0, grade: false, solution: true -->
```python
def appartient(T,n) : 
    while n not in T:
        return False
    return True
```

<!-- grade_id: cell-21706b32c82e7144, points: 2, grade: true, solution: false -->
```python
assert appartient([1,2,3,4],5) == False
assert appartient([1,2,3,4],1) == True
assert appartient([1],1) == True
assert appartient([1],2) == False
assert appartient([1,3,4,5,8,2],2) == True
```

## Etude du coût en temps.
Nous allons maintenant étudier la complexité de cet algorithme. Pour pouvoir comparer les temps d'éxécution, nous allons étudier le cas le plus dévaforable en temps.

**A quelle condition l'algorithme précédemment écrit prend-il le plus de temps d'exécution ?**

**Comment cela se traduit-il dans la fonction `compte_temps(n)` définie ci dessous ?**

<!-- grade_id: cell-7364916493cc3d43, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python
from time import process_time
from random import randint
    
def compte_temps(n):
    L = [randint(1,100) for i in range(n)]
    clef = 101
    t1_start = process_time()  
    appartient(L,clef)
    t1_stop = process_time() 
    return (t1_stop-t1_start)
```

```python
import matplotlib.pyplot as plt

x = [i for i in range(0,100000,5000)]
y = []
for i in x :
    y.append(compte_temps(i))
plt.plot(x,y)
```

Output:
```
[<matplotlib.lines.Line2D at 0x7f0d601df5c0>]
```

## Comparaison du coût en temps d'opérations élémentaires.
Considérons la fonction `appartient2(T,n)` définie ci dessous, 

```python
def appartient2(T,n): 
    for i in range(len(T)) :
        if T[i] == n :
            return True
    return False
    
def compte_temps2(n):
    L = [randint(1,100) for i in range(n)]
    clef = 101
    t1_start = process_time()  
    appartient2(L,clef)  
    t1_stop = process_time() 
    return (t1_stop-t1_start)
```

Et comparons maintenant l'exécution des deux fonctions :

```python
x = [i for i in range(0,100000,5000)]
y1,y2 = [],[]
for i in x :
    y1.append(compte_temps(i))
    y2.append(compte_temps2(i))
plt.plot(x,y1,label='boucle while')
plt.plot(x,y2,label='boucle for')
plt.legend()
```

**Proposer une explication à la différence de temps observé entre les deux mises en oeuvre de l'algorithme d'appartenance d'une valeur à un tableau**

<!-- grade_id: cell-3189db8a743b9f57, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Vérification de liste triée
## Fonction de vérification
> Ecrire une fonction `test_tri(T)` qui teste si un tableau d'entiers et trié par ordre croissant

<!-- grade_id: cell-0bdadc0e4a78f05c, points: 0, grade: false, solution: true -->
```python
def test_tri(T) :
    """ Test si une liste d'entiers est triè par ordre croissant
    Prec : Une liste d'entiers positifs ou nul, non vide
    Post : Un booléen, renvoie True si la liste est triée, False sinon"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fd1a1112c3dba871, points: 2, grade: true, solution: false -->
```python
assert test_tri([1,2,3,4,5,6]) == True
assert test_tri([1]) == True
assert test_tri([1,1,1,1]) == True
assert test_tri([4,3,2,1]) == False
assert test_tri([1,2,3,2,1]) == False
```

## Coût en temps
**On considére la fonction `compte_temps4(n)` définie ci-dessous, expliquer pourquoi cette fonction permet de trouver le coût maximal en temps de la fonction `test_tri(T)` ?**

<!-- grade_id: cell-32c8a35b57f5b09c, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
def compte_temps4(n):
    L = [i for i in range(n)]
    t1_start = process_time()  
    test_tri(L) 
    t1_stop = process_time() 
    return (t1_stop-t1_start)
```

```python
x = [i for i in range(0,500000,20000)]
y = []
for i in x :
    y.append(compte_temps4(i))
plt.plot(x,y)
```

**En déduire la compléxité de cette fonction de test de liste triée.**

<!-- grade_id: cell-0d87b831169a8a0b, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**On peut également comparer le coût en temps avec une fonction intégrée au langage.**

```python
def test_tri2(L) :
    if L==sorted(L) :
        return True
    else :
        return False
```

```python
def compte_temps5(n):
    L = [i for i in range(n)]
    t1_start = process_time()  
    test_tri2(L) 
    t1_stop = process_time() 
    return (t1_stop-t1_start)

x = [i for i in range(0,1000000,50000)]
y1,y2 = [],[]
for i in x :
    y1.append(compte_temps4(i))
    y2.append(compte_temps5(i))
    
plt.plot(x,y1,label = 'boucle for')
plt.plot(x,y2,label = 'fonction sorted')
```

**On remarque que la fonction `sorted` est plus rapide, mais cela n'enléve rien à son coût linéaire**